# SLU15 | Debugging. Exercise notebook

***

Now, we're going to solve a couple of bugs ourselves.

Use the suggestions from the learning notebook, your own logic, and Google (if required).

**WARNING**: 

We commented some cells out here. The reason for this is that they produce exceptions (and this behavior is expected), but our internal system doesn't allow exercise notebooks to contain cells that produce exceptions.

Uncomment these cells when you want to execute them, and comment them back out when you submit your solution.

## Start by importing these packages

```python
import hashlib

def _hash(s):
    return hashlib.blake2b(
        bytes(str(s), encoding='utf8'),
        digest_size=5
    ).hexdigest()
```

Our main task is to write a program that reads a text file that contains a list of numbers, calculates a "moving average" of that list, and writes the results to another text file named `result.txt`.

What we mean by a "moving average" is clear from the following example. If we want to calculate the averages of a list of numbers `[1, 2, 6, 3]`, we should begin with the first `[1]`, and the average of that number by itself is `1`. Then, we include the next number and calculate the average of `[1, 2]`, which is `1.5`. We add the next number, and the average of `[1, 2, 6]` becomes `3`. Finally, with the last number in the list, and the average of `[1, 2, 6, 3]` is `3`. The result of these moving averages, if we put them in another list, is `[1, 1.5, 3, 3]`.

Going back to our task, our program should create a final file named `result.txt` that contains all of our moving averages, each printed on a new line.

To accomplish this, we'll begin with a `nums` list, which is a list of numbers, and we'll use an element named `i` to store our averages. Our comparisons will continue until we hit the end of the index of `i`, so we need to be careful to include that last number in the list.

Let's make things even more interesting and force ourselves to implement the moving average calcuations ourselves, so we shouldn't use functions like len(), sum(), mean() for our calculations. It's okay to use these functions in any other situations, or course, but we want to test ourselves here. 

If you uncomment and try to execute the program below, you'll get an error:

`TypeError: unsupported operand type(s) for +=: 'int' and 'str'`.


```python
# with open('document.txt', 'r') as f:
#     numbers = f.readlines()
# n_items = 0
# current_sum = 0
# means = []
# for i in numbers:
#     n_items += 1
#     current_sum += i
#     means.append(current_sum / n_items)
# with open('result.txt', 'w') as f:
#     for num in means:
#         f.write(num)
```

This error already suggests where to look for issues, but let's try a different approach.

Let's divide our code into functions, so that it's easier to find bugs. That will also allow us to easily modify the program in the next steps.

### Task 1:
Create functions `read_nums`, `calculate_moving_avg`, `write_nums` from the code above.

Function `read_nums` needs to take a string `file_name` as a parameter and return the result of reading this file as a list of values.

Function `calculate_moving_avg` needs to take a list of integers `nums` as its parameter and return a list of moving averages as a result.

Function `write_nums` needs to take a string `file_name` and a list of floats `nums` as parameters and write each number to the file. This function shouldn't return anything.

You don't need to change anything in the code yet. Just divide the code into functions.

<!-- grade_id: cell-c806f18640ab4195, points: 0, grade: false, solution: true -->
```python
# def read_nums(file_name):
    # ...

# def calculate_moving_avg(nums):
    # ...

# def write_nums(file_name, nums):
#     ...


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c69436878e0ec76e, points: 2, grade: true, solution: false -->
```python
nums = read_nums('document.txt')
assert isinstance(nums, list), "Are you sure your read_nums function returns a list?"
assert len(nums) == 19, "Wrong number of items in the list returned by read_nums function"
assert _hash(str(nums)) == "0de4458d18", "The result of reading the file doesn't match the expected output"
```



As you can see, organizing our code into functions has made it much easier to understand what's wrong with it.

If you execute the code below, you should see the following error:

```
---------------------------------------------------------------------------
TypeError                                 Traceback (most recent call last)
<ipython-input-26-2c8baf7cfac4> in <module>
      1 nums = read_nums('document.txt')
----> 2 means = calculate_moving_avg(nums)
      3 write_nums('result.txt', means)

<ipython-input-25-bb04f4b79a0c> in calculate_moving_avg(nums)
     10     for i in numbers:
     11         n_items += 1
---> 12         current_sum += i
     13         means.append(current_sum / n_items)
     14     return means

TypeError: unsupported operand type(s) for +=: 'int' and 'str'
```

```python
 # nums = read_nums('document.txt')
 # means = calculate_moving_avg(nums)
 # write_nums('result.txt', means)
```



### Task 2:

It's important to read the errors, so pay attention to the error above and use it to answer the questions below.

Your answers will assign the letter of the correct option to the `answer_i` variable, where `i` is the question number.

For example, `answer_1="A"` assigns option `A` to question `1`. 

Some questions may have more than one correct answer, so, in that case, include all of their letters without spaces: 

`answer_i = "ABC"`

1. What type of error is the Exception above? 

```
A. IndexError
B. TypeError
C. NameError
D. SyntaxError
```

2. What's the error message.

```
A. calculate_moving_avg(nums)
B. Traceback (most recent call last)
C. unsupported operand type(s) for +=: 'int' and 'str'
D. TypeError
```

3. Which line or lines of code cause the errors?

```
A. current_sum += i 
B. means.append(current_sum / n_items)
C. write_nums('result.txt', means)
D. means = calculate_moving_avg(nums)
```

<!-- grade_id: cell-568fd9fe92bd054b, points: 0, grade: false, solution: true -->
```python
# answer_1 = ""
# answer_2 = ""
# answer_3 = ""

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2038524ce338dfa9, points: 6, grade: true, solution: false -->
```python
assert _hash(''.join(sorted(answer_1.upper()))) == '90a8cbb480', "Wrong exception type"
assert _hash(''.join(sorted(answer_2.upper()))) == 'f72385315c', "Wrong error message"
assert _hash(''.join(sorted(answer_3.upper()))) == 'd685f0ed6b', "Wrong error code lines"
```

As a reminder, our error is `TypeError: unsupported operand type(s) for +=: 'int' and 'str'`, and it's the result of calling `current_sum += i` in `calculate_moving_avg`.

What does that mean?

Well, it doesn't mean that the `calculate_moving_avg` function contains bugs. 

Rather, something wrong happened while executing `current_sum += i` inside `calculate_moving_avg`. Our mistake was to add an integer to a string in that line, but Python doesn't allow such operations.

Let's figure out why a string showed up there, because our plan was to add an integer to an integer, right?



### Task 3:




Your next task is to go through the code, line by line, analyze what happens with `current_sum` and `i`, and figure out why that string shows up.

Modify the created function `read_nums` in such a way that both `current_sum` and `i` can only be integers, which would make this error disappear.

**NOTE:**

There are other bugs in our code, but we'll leave those for other tasks later on. 

Remember. Only edit what is relevant to the exception we're dealing with right now.

<!-- grade_id: cell-949aca96afb49149, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-69d3c58b504eb6f1, points: 3, grade: true, solution: false -->
```python
nums = read_nums('document.txt')
means = calculate_moving_avg(nums)
assert isinstance(means, list), "calculate_moving_avg function is supposed to return a list"
assert len(means) == 19, "wrong number of items in the list returned by calculate_moving_avg function"
for num in means:
    assert isinstance(num, float), "Wrong value type. calculate_moving_avg function is supposed to return a list of floats"
```

If you solved the previous exercise correctly, our exception should have dissappeared, but there should be another exception if you execute the code below:

```
---------------------------------------------------------------------------
TypeError                                 Traceback (most recent call last)
<ipython-input-27-2c8baf7cfac4> in <module>
      1 nums = read_nums('document.txt')
      2 means = calculate_moving_avg(nums)
----> 3 write_nums('result.txt', means)

<ipython-input-25-8ffcbd828a96> in write_nums(file_name, nums)
     23     with open(file_name, 'w') as f:
     24         for num in nums:
---> 25             f.write(num)
     26 
     27 

TypeError: write() argument must be str, not float
```


```python
# nums = read_nums('document.txt')
# means = calculate_moving_avg(nums)
# write_nums('result.txt', means)
```



# Task 4:


Just like we did at the start of task 2, answer the following questions by assigning the letter of the correct option to the `answer_i` variable, where `i` is the question number.

For example, `answer_1="A"` assigns option `A` to question `1`. 

Some questions may have more than one correct answer, so, in that case, include all of their letters without spaces: 

`answer_i = "ABC"` assigns the letters `A`, `B`, and `C` to question `1`.

1. What's the type of the new Exception?

```
A. IndexError
B. TypeError
C. NameError
D. SyntaxError
```

2. What's its error message.

```
A. write() argument must be str, not float
B. Traceback (most recent call last)
C. f.write(num)
D. write_nums('result.txt', means)
```

3. Which line or lines of code cause the errors?

```
A. current_sum += i 
B. means.append(current_sum / n_items)
C. write_nums('result.txt', means)
D. f.write(num)
```

<!-- grade_id: cell-9e13dd102b96d202, points: 0, grade: false, solution: true -->
```python
# answer_1 = ""
# answer_2 = ""
# answer_3 = ""

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-a2a266bf76e087e1, points: 6, grade: true, solution: false -->
```python
assert _hash(''.join(sorted(answer_1.upper()))) == '90a8cbb480', "Wrong exception type"
assert _hash(''.join(sorted(answer_2.upper()))) == 'b685b2f862', "Wrong error message"
assert _hash(''.join(sorted(answer_3.upper()))) == '94c05104ae', "Wrong error code lines"
```

## Task 5 - Part 1:

Once again, as a reminder, our current error is `TypeError: write() argument must be str, not float`, and it's the result of calling `f.write(num)` in `write_nums` function.

What does that mean?

It doesn't mean that the `write_nums` function contains bugs. 

However, there was something wrong while executing `f.write(num)` inside of `calculate_moving_avg`. Our mistake was expecting that expression to accept a float-type number. 

If we check the `write` method documentation [here](https://www.w3schools.com/python/ref_file_write.asp), we find the following:

`parameter: byte --- The text or byte object that will be inserted.`

This means that the `write` method only allows strings and byte objects, not floats.

Seems like `num` is a float number.


Your task is to modify the code in such a way so that the `write` method works.

**NOTE:** 

There might still be some other bugs in our code, but we'll deal with those later in other tasks. Remember not to edit anything that's not related to our current exception.

<!-- grade_id: cell-082945ee7566cb3c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

```python
nums = read_nums('document.txt')
means = calculate_moving_avg(nums)
write_nums('result.txt', means)
```

## Task 5 - Part 2:

If you did everything correctly in the previous exercise, there should be no more exceptions.

But it doesn't mean that our program works 100% correctly, right?

We need to validate that the results are correct.

Let's read our main task once again.

```

Our main task is to write a program that reads a text file that contains a list of numbers, calculates a "moving average" of that list, and writes the results to another text file named `result.txt`.

...our program should create a final file named `result.txt` that contains all of our moving averages, each printed on a new line.

```

Did we meet the requirements? If you opened the `result.txt` file, you'd find out that it looks like this:

<img src="./assets/result.png" width="900">

And the task description says that the results file should contain numbers placed on new lines.

Your new task is to modify the code in such a way, so that each number is on a new line.
The last line should be a number, not an empty string.

e.g., correct:

`10.5`

`3.5`

`1.2`


wrong:

`10.5`

`3.5`

`1.2`

` `

<!-- grade_id: cell-ed7689c30aba19db, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-c9e2a417b3135004, points: 3, grade: true, solution: false -->
```python
nums = read_nums('document.txt')
means = calculate_moving_avg(nums)
write_nums('result.txt', means)
with open('result.txt', 'r') as f:
    content = f.readlines()
assert len(content) == 19, "Are you sure there is the same number of lines in the text file, as in the moving averages list?"
for i in range(len(content)):
    line = content[i]
    if i != len(content) - 1:
        assert line[-1:] == "\n", "Are you sure there is a newline symbol after each number (besides the last one)?"
    else:
        assert line[-1:] != "\n", "There should be no newline symbol at the end of the file"
    num = float(line.strip()) # if line.strip() can be converted to float, everything is alright
```

### And now... Mission Accomplished!!! 🎉🎉🎉
Now you can rest for this week, because your work was amaaaaaaazing (Again)!

**But don't forget...**

***

# Submit your work!

To grade your exercise notebook and submit your work to the portal, [follow the instructions in the weekly workflow!](https://github.com/LDSSA/ds-prep-course-2024/blob/main/weekly-workflow.md#link-to-grading)